# Question 1

This question will focus entirely on NumPy arrays: vectorized programming, slicing, and broadcasting.

### Part A

In this question, you'll implement the vector dot product.

Write a function which:

 - is named `dot`
 - takes two NumPy arrays as arguments
 - returns one number: the floating-point dot product of the two input vectors

Recall how a [dot product](https://en.wikipedia.org/wiki/Dot_product) works: corresponding elements of two arrays are multiplied together, then all these products are summed.

For example: if I have two NumPy arrays `[1, 2, 3]` and `[4, 5, 6]`, their dot product would be `(1*4) + (2*5) + (3*6)`, or `4 + 10 + 18`, or 32.

You can use NumPy arrays, and the `np.sum()` function, but no other NumPy functions. Also, **no loops are allowed!** This has to be done using *vectorized* computation.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
import numpy as np
def dot(arr1, arr2):
    newArr = arr1*arr2
    return np.sum(newArr)
    
```

<!-- grade_id: q1a_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(57442)

x1 = np.random.random(10)
x2 = np.random.random(10)
np.testing.assert_allclose(x1.dot(x2), dot(x1, x2), rtol = 1e-5)
```

<!-- grade_id: q1a_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(495835)

x1 = np.random.random(100)
x2 = np.random.random(100)
np.testing.assert_allclose(x1.dot(x2), dot(x1, x2), rtol = 1e-5)
```

### Part B

Write a function which:

 - is named `less_than`
 - takes two arguments: a NumPy array, and a floating-point number
 - returns a NumPy array
 
You should use a boolean mask to return only the values in the NumPy array that are less than the specified floating-point value (the second parameter). No loops are allowed, or any built-in functions or loops.

For example, `less_than([1, 2, 3], 2.5)` should return a NumPy array of `[1, 2]`.

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python
def less_than(arr,num):
    boolArr = arr < num
    newArr = arr[boolArr]
    return newArr
    
```

<!-- grade_id: q1b_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(85928)

x = np.random.random((10, 20, 30))
t = 0.001
y = np.array([ 0.0005339 ,  0.00085714,  0.00091265,  0.00037283])
np.testing.assert_allclose(y, less_than(x, t), rtol = 1e-4)
```

<!-- grade_id: q1b_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(8643)

x2 = np.random.random((100, 100, 10))
t2 = 0.0001
y2 = np.array([  2.91560413e-06,   6.80065620e-06,   3.63294064e-05,
         7.50659065e-05,   1.61602031e-06,   9.37205052e-05])
np.testing.assert_allclose(y2, less_than(x2, t2), rtol = 1e-05)
```

### Part C

Write a function which:

 - is named `greater_than`
 - takes two arguments: a NumPy array, and a threshold number (float)
 - returns a NumPy array
 
You should use a boolean mask to return only the values in the NumPy array that are greater than the specified `threshold` value (the second parameter). No loops are allowed, or built-in functions, or NumPy functions.

For example, `greater_than([1, 2, 3], 2.5)` should return a NumPy array of `[3]`.

<!-- grade_id: q1c, points: 0, grade: false, solution: true -->
```python
def greater_than(arr,num):
    boolArr = arr > num
    newArr = arr[boolArr]
    return newArr
    
```

<!-- grade_id: q1c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(592582)

x = np.random.random((10, 20, 30))
t = 0.999
y = np.array([ 0.99910167,  0.99982779,  0.99982253,  0.9991043 ])
np.testing.assert_allclose(y, greater_than(x, t), rtol = 1e-5)
```

<!-- grade_id: q1c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(689388)

x2 = np.random.random((100, 100, 10))
t2 = 0.9999
y2 = np.array([ 0.99997265,  0.99991169,  0.99998906,  0.99999012,  0.99992325,
        0.99993289,  0.99996637,  0.99996416,  0.99992627,  0.99994388,
        0.99993102,  0.99997486,  0.99992968,  0.99997598])
np.testing.assert_allclose(y2, greater_than(x2, t2), rtol = 1e-05)
```

### Part D

Write a function which:

 - is named `in_between`
 - takes three parameters: a NumPy array, a lower threshold (float), and an upper threshold (float)
 - returns a NumPy array
 
You should use a boolean mask to return only the values in the NumPy array that are *in between* the two specified threshold values, `lower` and `upper`. No loops are allowed, or built-in functions, or NumPy functions.

For example, `in_between([1, 2, 3], 1, 3)` should return a NumPy array of `[2]`.

Hint: you can use your functions from Parts C and D to help!

```python
a = np.array([1,2,3])
b=1
c=3

boolArr = (a > b) & (a < c)
print(boolArr)    
```

Output:
```
[False  True False]
```

<!-- grade_id: q1d, points: 0, grade: false, solution: true -->
```python
def in_between(arr,low_num,upp_num):
    boolArr = (arr > low_num) & (arr < upp_num)
    newArr = arr[boolArr]
    return newArr
```

<!-- grade_id: q1d_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(7472)

x = np.random.random((10, 20, 30))
lo = 0.499
hi = 0.501
y = np.array([ 0.50019884,  0.50039172,  0.500711  ,  0.49983418,  0.49942259,
        0.4994417 ,  0.49979261,  0.50029046,  0.5008376 ,  0.49985266,
        0.50015914,  0.50068227,  0.50060399,  0.49968918,  0.50091042,
        0.50063015,  0.50050032])
np.testing.assert_allclose(y, in_between(x, lo, hi), rtol = 1e-6)
```

<!-- grade_id: q1d_test2, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(14985)

x = np.random.random((30, 40, 50))
lo = 0.49999
hi = 0.50001
y = np.array([ 0.50000714,  0.49999045])
np.testing.assert_allclose(y, in_between(x, lo, hi), rtol = 1e-6)
```

### Part E

Write a function which:

 - is named `not_in_between`
 - takes three parameters: a NumPy array, a lower threshold (float), and an upper threshold (float)
 - returns a NumPy array
 
You should use a boolean mask to return only the values in the NumPy array that are *NOT in between* the two specified threshold values, `lower` and `upper`. No loops are allowed, or built-in functions, or NumPy functions.

For example, `not_in_between([1, 2, 3, 4], 1, 3)` should return a NumPy array of `[4]`.

Hint: you can use your functions from Parts C and D to help!

```python
a = np.array([1,2,3,4])
b=1
c=3

boolArr = (a >= b) & (a <= c)
#switchArr = not boolArr
print(~boolArr)    
```

Output:
```
[False False False  True]
```

<!-- grade_id: q1e, points: 0, grade: false, solution: true -->
```python
def not_in_between(arr,low_num,upp_num):
    boolArr = (arr >= low_num) & (arr <= upp_num)
    newArr = arr[~boolArr]
        
    return newArr
```

<!-- grade_id: q1e_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(475185)

x = np.random.random((10, 20, 30))
lo = 0.001
hi = 0.999
y = np.array([  9.52511605e-04,   8.62993716e-04,   3.70243252e-04,
         9.99945849e-01,   7.21751759e-04,   9.36931041e-04,
         5.10792605e-04,   6.44911672e-04])
np.testing.assert_allclose(y, not_in_between(x, lo, hi), rtol = 1e-6)
```

<!-- grade_id: q1e_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(51954)

x = np.random.random((30, 40, 50))
lo = 0.00001
hi = 0.99999
y = np.array([  8.46159001e-06,   9.99998669e-01,   9.99993873e-01,
         5.58488698e-06,   9.99993348e-01])
np.testing.assert_allclose(y, not_in_between(x, lo, hi), rtol = 1e-6)
```

### Part F

Write a function which:

 - is named `reverse_array`
 - takes 1 parameter: a 1D NumPy array of data
 - returns the 1D NumPy array, reversed

This function uses fancy indexing to reverse the ordering of the elements in the input array, and returns the reversed array. You *cannot* use the `[::-1]` notation, nor the built-in `reversed` method, or any other Python function or loops. You **can** use the `list()`, `range()`, and `np.arange()` functions, however, and only some or all of those (but again, no loops!).

Hint: Construct a list of indices and use NumPy fancy indexing to reverse the ordering of the elements in the input list, then return the reversed array.

<!-- grade_id: q1f, points: 0, grade: false, solution: true -->
```python
#This doesn't work
def reverse_array(arr):
    ran= range(arr)
    length = len(arr)
    newList = (arr[length-1:], arr[])
    start = arr[length -1:]
    max_index = length -1
    iteration = length / 2
    return newList
    
```

Output:
```
File "<ipython-input-35-8bd8c3cadfc1>", line 5
    newList = (arr[length-1:], arr[])
                                   ^
SyntaxError: invalid syntax
```

```python
def reverse_array(array):
    start = array.shape[0]-1
    end = -1
    move = -1
    index = np.arange(start, end, move)
    print(index)
    newList = array[index]
    return newList
```

<!-- grade_id: q1f_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(5748)

x1 = np.random.random(75)
y1 = x1[::-1]  # Sorry, you're not allowed to do this!
np.testing.assert_allclose(y1, reverse_array(x1), rtol = 1e-6)
```

Output:
```
[74 73 72 71 70 69 68 67 66 65 64 63 62 61 60 59 58 57 56 55 54 53 52 51
 50 49 48 47 46 45 44 43 42 41 40 39 38 37 36 35 34 33 32 31 30 29 28 27
 26 25 24 23 22 21 20 19 18 17 16 15 14 13 12 11 10  9  8  7  6  5  4  3
  2  1  0]
```

<!-- grade_id: q1f_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(68382)

x2 = np.random.random(581)
y2 = x2[::-1]  # Sorry, you're not allowed to do this!
np.testing.assert_allclose(y2, reverse_array(x2), rtol = 1e-6)
```

Output:
```
[580 579 578 577 576 575 574 573 572 571 570 569 568 567 566 565 564 563
 562 561 560 559 558 557 556 555 554 553 552 551 550 549 548 547 546 545
 544 543 542 541 540 539 538 537 536 535 534 533 532 531 530 529 528 527
 526 525 524 523 522 521 520 519 518 517 516 515 514 513 512 511 510 509
 508 507 506 505 504 503 502 501 500 499 498 497 496 495 494 493 492 491
 490 489 488 487 486 485 484 483 482 481 480 479 478 477 476 475 474 473
 472 471 470 469 468 467 466 465 464 463 462 461 460 459 458 457 456 455
 454 453 452 451 450 449 448 447 446 445 444 443 442 441 440 439 438 437
 436 435 434 433 432 431 430 429 428 427 426 425 424 423 422 421 420 419
 418 417 416 415 414 413 412 411 410 409 408 407 406 405 404 403 402 401
 400 399 398 397 396 395 394 393 392 391 390 389 388 387 386 385 384 383
 382 381 380 379 378 377 376 375 374 373 372 371 370 369 368 367 366 365
 364 363 362 361 360 359 358 357 356 355 354 353 352 351 350 349 348 347
 346 345 344 343 342 341 340 339 338 337 336 335 33
```

```python

```

```python

```